# Wstęp

Na tych zajęciach będziemy korzystali z Jupytera/Colaboratory do pisania kodu i oddawania prac domowych.

Każdy *notebook* Jupytera/Colaboratory składa się z wielu komórek (*cell*). Na przykład, ten tekst jest w komórce typu tekstowego (*Markdown cell* w Jupyterze, *TEXT* w  Colaboratory). Poniższa komórka zawiera kod (jest typu *code*).

```python
# to jest komórka z kodem
2+2
```

Kliknij na komórkę, by ją wybrać. 

W Jupyterze możesz dowiedzieć się jakiego typu jest komórka poprzez kliknięcie na nią i spojrzenie na pasek narzędzi. Gdy klikniesz na tę komórkę, zobaczysz typ *Markdown* jak na obrazku:

![](https://www.mimuw.edu.pl/~mm319369/priv/d73890416bec03ff3e2b3756af8c941c/images/cell-type-menu.png)

---

## Komórki z kodem

Komórki z kodem mogą zawierać dowolny poprawny kod Pythona. Gdy uruchomisz komórkę, kod jest wykonywany, a wyjście programu wyświetlane poniżej.

<div class="alert alert-info">
Można wykonać zawartość zaznaczonej komórki na wiele sposobów:
<ul>
<li> <b><code>Ctrl+Enter/Cmd+Enter</code></b> (komórka pozostanie zaznaczona), </li>
<li> <b><code>Shift+Enter</code></b> (zaznaczy kolejną komórkę), </li>
<li> <b><code>Alt+Enter</code></b> (stworzy komórkę poniżej), </li>
<li> klikając guzik **play** znajdujący się:
    <ul>
    <li> na pasku narzędzi w Jupyterze (obok kwadratu), </li>
    <li> na lewej belce koło kodu w Colaboratory. </li>
    </ul></li>
</ul>
</div>

<b>Uwaga!</b> W Colaboratory musimy poczekać na przydzielenie nam zasobów. Status znajduje się w prawym górnym rogu.

Spróbuj wykonać poniższą komórkę i zobacz, co wypisze:

```python
print("Printing cumulative sum from 1-10:")
total = 0
for i in range(1, 11):
    total += i
    print("Sum of 1 to " + str(i) + " is: " + str(total))
print("Done printing numbers.")
```

Zauważ, że poniżej komórki znalazły się rezultaty wykonania instrukcji `print`.
Kolejny przykład wypisuje tylko ostateczną sumę:

```python
total = 0
for i in range(1, 11):
    total += i
print(total)
```

Innym sposobem na wypisanie czegoś jest po prostu pozostawienie tego obiektu w ostatniej linijce. Możemy zmienić powyższy przykład na:

```python
total = 0
for i in range(1, 11):
    total += i
total
```

Jednak, to *nie zadziała* jeśli niżej będą jeszcze jakieś inne linie. Na przykład, po dodaniu na końcu dodatkowej instrukcji `print`, nie wypisze się zawartość zmiennej `total` tak, jakbyśmy tego oczekiwali:

```python
total = 0
for i in range(1, 11):
    total += i
total
print("Done computing total.")
```

Jeśli zazwyczaj używasz Pythona 2, zwróć uwagę na obowiązkowe nawiasy wokół funkcji `print` w Pythonie 3.

---

# Komórki tekstowe

To jest przykładowa *komórka tekstowa*.
Komórki tekstowe możesz edytować po **dwukrotnym kliknięciu** ich. Spróbuj kliknąć w tę.

Krótki przewodnik po składni znajdziecie [tutaj](https://docs.microsoft.com/pl-pl/contribute/help-crr/help-content/contribute/contribute-how-to-write-use-markdown).
Rozszerzony przykład w Colaboratory znajdziecie w tym [notebooku](/notebooks/markdown_guide.ipynb).

Możecie również pisać równania matematyczne w składni LaTeX-a: $\sum_i^\infty \sqrt{i+2}$.

Aby zamienić tę komórkę na sformatowaną, w Jupyterze trzeba ją uruchomić, a w Colaboratory wystarczy kliknąć gdzie indziej.

# Dodawanie i przesuwanie komórek, zapisywanie

Nowe komórki można tworzyć za pomocą:

W Jupyterze:
1. klikając w ikonę **plus**a (na pasku narzędiowym), który doda komórkę pod obecnie zaznaczoną,
2. typ możemy zmienić za pomocą menu (domyślnie *Code*).

W Colaboratory:
* klikając w **+ CODE** lub w **+ TEXT** na pasku narzędziowym,
* analogiczne guziki pojawiają się po najechaniu na granicę między komórkami.

Komórki można przesuwać za pomocą **strzałek góra/dół** w pasku narzędziowym. Można zaznaczyć wiele przez kliknięcie na jedną, a następnie, przytrzymując **Shift**, kliknięcie na następną. W Colaboratory możemy użyć "lassa" gdy zaczniemy zaznaczanie poza regionem komórek.

Zaznaczone komórki możemy wycinać, kopiować i wklejać za pomocą odpowiedniego menu. W Colaboratory, ponadto, istnieje wygodny przycisk śmietnika w rogu komórki.

**Uwaga!**
Pracując w Colaboratory, zmiany w notebooku zapisywane są automatycznie. W Jupyterze nie zawsze tak jest i czasem należy kliknąć w ikonę zapisywania na pasku narzędzi.

# Praca z Pythonem

Gdy tworzysz i uruchamiasz nowy notebook, uruchamiasz również system odpowiedzialny za intepretację kodu. W Jupyterze nazywa się to *kernel*em, a w Colaboratory *runtime*'em. Ten system działa w tle: gdy uruchamiasz komórkę z kodem, jej treść wykonywana jest przez ten system, a wyjście przekazywane z powrotem do *notebook*a.

Działa to podobnie to interpretera Pythona w konsoli: np. musisz zadeklarować zmienne, zanim ich użyjesz. Zobacz, co stanie się, gdy wykonasz poniższy kod:

```python
a
```

Jak widzisz, wyjątki są ładnie formatowane. Powyżej mamy problem z tym, że zmienna `a` nie istnieje. Zmień powyższą komórkę tak, aby zdefiniować najpierw zmienną `a` (np. przypisz do niej 1). Gdy to zrobisz, możesz wykonać powyższą komórkę bez problemów.

```python
print("The value of 'a' is: " + str(a))
```

Wykonanie powyższej komórki powinno już działać, gdyż zmienna `a` została już zdefiniowana. Zmienne zostają w środowisku działającym w tle. Użyj polecenia `%whos`, by zobaczyć jakie zmienne sa dostępne:

```python
%whos
```

Jeśli przykłady z sumowaniem zostały wykonane, zobaczysz także `total` i `i`. Dzieje się tak, gdyż każda wykonana komórka zmienia stan systemu.

Zauważ, że polecania zaczynające się od procenta (%) albo dwóch (%%) są specjalnymi poleceniami systemu działającego w tle (IPython). Nazywa się je *magic*s i działają tylko z nim. Innym, specjalnym modyfikatorem jest wykrzyknik (!), który wykonuje kod w terminalu:

```python
!echo 22
!ls || dir
```

## Resetowanie stanu

Dobrą praktyką przy pracy z Jupyterem jest, aby komórki wykonane od góry do dołu w świeżym środowisku działały zgodnie z oczekiwaniem. (W szczególności, **będzie to wymagane przy oddawaniu prac domowych**.) Częstym błędem jest używanie w komórkach zmiennych, które kiedyś zostały zdefiniowane, ale później kod je definiujący został skasowany.

Aby sprawdzić, czy tak jest, zresetuj najpierw stan środowiska poprzez wybranie w menu Kernel/Runtime -> Restart.

W Jupyterze można też klikąć ikonkę na pasku narzędziowym:
![](https://www.mimuw.edu.pl/~mm319369/priv/d73890416bec03ff3e2b3756af8c941c/images/restart-kernel-button.png)

Następnie, wykonaj wszystkie komórki przez wybranie w menu:
* w Jupyterze: **Cell$\rightarrow$Run All**,
* w Colaboratory: **Runtime$\rightarrow$Run All**,

Stan środowiska ze zmiennymi nie jest nigdzie zapisywany. W Colaboratory nasz stan zniknie po około 12 godzinach. Podobnie stanie się, gdy np. wyłączymy komputer z naszym własnym Jupyterem.

## Przerywanie wykonania

Jeśli dana komórka działa zbyt długo, można przerwać jej wykonanie. Wykonaj poniższą komórkę i wybierz z menu **Kernel/Runtime -> Interrupt (execution)** zanim się skończy. W Jupyterze możesz też kliknąć kwadrat na pasku narzędzi. (Pod Windowsem i w Colaboratory może to czasem nie działać.)

```python
total = 0
for i in range(50000000):
    total += i*i # too long, interrupt me!
print("Done computing total:", total)
```

## Dopełnianie TAB-em

Jupyter/Colaboratory wspiera dopełnianie po wciśnięciu klawisza **Tab**. Wykonaj najpierw poniższą komórkę. Następnie, wciśnij **Tab** w pozostałych, po umieszczeniu kursora na ich końcu:

```python
import math
```

```python
t
```

```python
math.
```

W Colaboratory poniższe wyświetli etykietę z dokumentacją:

```python
math.cos(
```

Wykonanie komórki ze znakiem zapytania na końcu (?) wyświetla dłuższą dokumentację w okienku poniżej. Można je zamknąć iksem w prawym górnym rogu.

```python
math.atan2?
```

## Zaawansowane wyjście

Jupyter pozwala wyświetlać nie tylko tekst, ale na przykład też wykresy. Uruchom poniższą komórkę, by go zobaczyć:

```python
import numpy as np
from matplotlib import pyplot as plt

x = np.linspace(0, 10, 100) # 100 liczb od 0 do 10, w równych odstępach
y = np.sin(x)

plt.plot(x, y)
plt.title("Wykres sinusa")
plt.show()
```

## Formularze w Colaboratory

W Colaboratory możemy tworzyć proste formularze, które zmieniają zmienne w środowisku. Uruchom poniższą komórkę. Przy niej powinien pojawić się slider z możliwością wyboru wartości. Po jego zmianie, komórka powinna automatycznie się wykonać.

```python
#@title Tytuł formularza { run: "auto", display-mode: "both" }
a = 0.1 #@param {type:"slider", min:0, max:2, step:0.1}
print(a+2)

# Tak w razie czego poniżej importujemy jeszcze raz
import numpy as np
from matplotlib import pyplot as plt

x = np.linspace(0, 10, 100)
y = np.sin(x) + a * np.sin(x*4)
plt.plot(x, y)
plt.show()
```

Aby utworzyć podobny formularz, zaznaczcie kursorem na poniższą komórkę, a następnie kliknijcie w **+form** w jej prawym górnym rogu. Następnie kliknijcie w **+field**, wybierzcie nazwę zmiennej (*Variable name*), zmieńcie typ (*Variable type*) na *number* oraz rodzaj pola (*Form field type*) na slider. Możecie tam wybrać zakres (*Min* - *Max*) oraz rozmiar kroku (*Step*). Kliknijcie w ikonkę edycji koło tytułu (*Default title text*) i zaznaczcie automatyczne uruchamianie komórki (*Auto-execute cell when fields change*)"

```python
print(a)
```

## Pobieranie plików z Colaboratory

W Colaboratory niestety nie mamy łatwego dostępu do stworzonych plików. Możemy je pobrać w następujący sposób:

```python
from google.colab import files

with open('example.txt', 'w') as f:
  f.write('some content')

files.download('example.txt')
```

## Testy

W zadaniach programistycznych otrzymacie częściowo wypełnione notebooki, które będą zawierały wiele komórek, które będziecie musieli uzupełnić. Niektóre komórki będą zawierały testy sprawdzające poprawność rozwiązania. Aby dostać punkty za dane testy, cała komórka musi wykonać się bez błędu. (Proszę nie zmieniać ich zawartości ani ich nie usuwać – to nie pomoże, a tylko utrudni nam pracę.)
Będą też pytania z oczekiwaną odpowiedzią tekstową w danej komórce.

Poniżej przykład schematu zadania:

<div class="alert alert-success">
Zaimplementuj funkcję <code>witaj</code> i upewnij się, że poniższe testy przechodzą bez problemu. Zacznij od skasowania linijki z `raise NotImplementedError`, napisz swoje rozwiązanie, a następnie uruchom ponownie tę komórkę i komórkę z testami.

Pamiętaj, że za każdym razem, gdy zmieniasz treść komórki (również kod funkcji), musisz ją ponownie uruchomić, aby pozostałe komórki mogły dostrzec zmianę!
</div>

<!-- grade_id: hello, points: 0, grade: false, solution: true -->
```python
def witaj(imie):
    """Zwraca wiadomość zawierającą "Witaj, <imie>!",
    gdzie <imie> zostało ponade jako argument.
    
    Argumenty
    ---------
    imie : string
        Imię osoby do przywitania
        
    Zwracana wartość
    ----------------
    wiadomość zawierająca "Witaj, <imie>!"
    
    """
    raise NotImplementedError
```

```python
# spróbuj wykonać tę komórkę z Twoim imieniem i zobacz, co zwróci
witaj("TWOJE IMIĘ TUTAJ")
```

```python
# Podpowiedż: jeśli kod funkcji wydaje się w porządku,
# jednak z jakiegoś powodu testy nie przechodzą,
# sprawdź czy zwracasz wartość za pomocą `return`,
# a nie wypisujesz jej przez `print`.

# Poniższy kod powinien wypisać tylko jedną linijkę.
# Wypisanie "Twoja funkcja zwróciła:  None"
# sugeruje błąd opisany powyżej

message = witaj("Janie Kowalski")
print("Twoja funkcja zwróciła: ", str(message))
```

<div class="alert alert-info">
Podczas rozwiązywania zadania mogą przydać Ci się Twoje własne testy. Możesz je napisać np. poniżej korzystając z takiej samej składni. Nie będą one uwzględniane przy ocenianiu.</div>

```python
# dodaj własne testy tutaj

```

<!-- grade_id: test_hello, points: 1, grade: true, solution: false -->
```python
"""Kod z testami. Ta komórka powinna zakończyć się wypisaniem "Success" aby ją zaliczyć."""
assert witaj("Świecie") == "Witaj, Świecie!"
assert witaj("świecie") == "Witaj, świecie!"
assert witaj("Janie Kowalski") == "Witaj, Janie Kowalski!"
assert witaj("123") == "Witaj, 123!"

print("Success!")
```

Powyższy opis powstał na podstawie przykładu użycia systemu do oceniania notebooków: [jhamrick/nbgrader-demo](https://github.com/jhamrick/nbgrader-demo/blob/05a5c90582fc3a575e70f94a0cc214c275311264/instructor/source/ps1/1%20-%20Introduction%20to%20the%20IPython%20notebook.ipynb) i [wstępu do Colaboratory](https://colab.research.google.com/notebooks/basic_features_overview.ipynb).